```python
# 請輸入姓名以及學號，以證名為本人撰寫
Name = ""
NID = ""
```

## 字串處理 (30 points)

請將下面 `Str()` 方法完成，輸入一段字串後，再輸入 **被取代的文字** 以及 **取代的文字**，並輸出三個結果，分別為：
1. 所有文字轉換為大寫的字串
2. 所有文字轉換為小寫的字串
3. 被取代文字後的字串
<br></br><br></br>
注意：使用 `upper()`、`lower()` 以及 `replace()` 實作

<!-- grade_id: str, points: 0, grade: false, solution: true -->
```python
def Str():
    """輸入目標字串、被取代的文字以及取代的文字，並印出題目格式"""""
    
    # YOUR CODE HERE
    first = input('')
    second = input('')
    third = input('')
    print(first.upper())
    print(first.lower())
    print(first.replace(second, third))
```

當輸入 `TA is handsome`、`is` 以及 `isn't` 時，`Str()` 方法應該印出 <br></br><br></br>
`TA IS HANDSOME`<br></br>
`ta is handsome`<br></br>
`TA isn't handsome`<br></br><br></br>
請執行以下程式碼確認結果。

```python
Str()
```

<!-- grade_id: correct_str, points: 30, grade: true, solution: false -->
```python
"""檢查 Str() 結果是否與以下測試案例相同"""
import unittest
from unittest import mock
from unittest.mock import patch
from io import StringIO

class TestStr(unittest.TestCase):
    
    input_str = "FCU is rich"
    old = 'is'
    new = 'isn\'t'
    
    @patch('builtins.input', side_effect=[input_str, old, new])
    def test_Str(self, mock_input):
        with mock.patch('sys.stdout', new = StringIO()) as output:
            Str()
            self.assertEqual(output.getvalue(), "FCU IS RICH\nfcu is rich\nFCU isn't rich\n")

suite = unittest.TestSuite()
suite.addTest(TestStr('test_Str'))
result = unittest.TextTestRunner(verbosity=2).run(suite)

assert result.wasSuccessful() == True
```