```python
# 請輸入姓名以及學號，以證名為本人撰寫
Name = ""
NID = ""
```

## 十二星座 (25 points)

請將下面 `constellation()` 方法完成，輸入月份以及日期並回傳 **星座**。<br></br><br></br>
十二星座如下：
* 水瓶座 ( 1/20 - 2/18 )：Aquarius
* 雙魚座 ( 2/19 - 3/20 )：Pisces
* 牧羊座 ( 3/21 - 4/20 )：Aries
* 金牛座 ( 4/21 - 5/20 )：Taurus
* 雙子座 ( 5/21 - 6/21 )：Gemini
* 巨蟹座 ( 6/22 - 7/22 )：Cancer
* 獅子座 ( 7/23 - 8/22 )：Leo
* 處女座 ( 8/23 - 9/22 )：Virgo
* 天秤座 ( 9/23 - 10/23 )：Libra
* 天蠍座 ( 10/24 - 11/22 )：Scorpio
* 射手座 ( 11/23 - 12/21 )：Sagittarius
* 摩羯座 ( 12/22 - 1/20 )：Capricorn

<!-- grade_id: constellation, points: 0, grade: false, solution: true -->
```python
def constellation():
    """輸入月份及日期，並印出為哪一個星座"""""
    
    # YOUR CODE HERE
    month = int(input())
    day = int(input())
    num = 0
    total = 0
    if month == 1:
        num = 31
    if month == 2:
        num = 59
    if month == 3:
        num = 90
    if month == 4:
        num = 120
    if month == 5:
        num = 151
    if month == 6:
        num = 181
    if month == 7:
        num = 212
    if month == 8:
        num = 243
    if month == 9:
        num = 273
    if month == 10:
        num = 304
    if month == 11:
        num = 334
    if month == 12:
        num = 365
        
    total = num + day
    if total<53:
        total = total + 365
    if total>=388 and total<=417:
        print('Capricorn')
    if total>=53 and total<=78:
        print('Aquarius')
    if total>=79 and total<=110:
        print('Pisces')
    if total>=111 and total<=140:
        print('Aries')
    if total>=141 and total<=171:
        print('Taurus')
    if total>=172 and total<=201:
        print('Gemini')
    if total>=202 and total<=234:
        print('Cancer')
    if total>=235 and total<=265:
        print('Leo')
    if total>=266 and total<=295:
        print('Virgo')
    if total>=296 and total<=326:
        print('Libra')
    if total>=327 and total<=356:
        print('Scorpio')
    if total>=357 and total<=387:
        print('Sagittarius')
```

當月份為 10 日期為 5 時，`constellation()` 方法應該印出 `Libra`，請執行以下程式碼確認結果。

```python
constellation()
```

<!-- grade_id: test_constellation, points: 25, grade: true, solution: false -->
```python
"""檢查 constellation() 結果是否與以下測試案例相同"""
import unittest
from unittest import mock
from unittest.mock import patch
from io import StringIO

class TestConstellation(unittest.TestCase):
    
    testcase_1 = [3, 7]
    testcase_2 = [5, 4]
    
    @patch('builtins.input', side_effect=[testcase_1[0], testcase_1[1]])
    def test_constellation_1(self, mock_input):
        with mock.patch('sys.stdout', new = StringIO()) as output:
            constellation()
            self.assertEqual(output.getvalue(), "Pisces\n")
            
    @patch('builtins.input', side_effect=[testcase_2[0], testcase_2[1]])
    def test_constellation_2(self, mock_input):
        with mock.patch('sys.stdout', new = StringIO()) as output:
            constellation()
            self.assertEqual(output.getvalue(), "Taurus\n")

suite = unittest.TestSuite()
suite.addTest(TestConstellation('test_constellation_1'))
suite.addTest(TestConstellation('test_constellation_2'))
result = unittest.TextTestRunner(verbosity=2).run(suite)

assert result.wasSuccessful() == True
```