```python
# 請輸入姓名以及學號，以證名為本人撰寫
Name = ""
NID = ""
```

## 計算圓的周長以及面積 (30 points)

請將下面 `circle()` 方法完成，輸入圓的 **直徑** 並計算圓周長以及面積。<br></br><br></br>
注意：直徑為 `float` 型態

<!-- grade_id: circle, points: 0, grade: false, solution: true -->
```python
def circle():
    """計算圓周長以及面積，輸入直徑，並印出圓周長以及面積"""""
    
    # YOUR CODE HERE
    pi = 3.14
    a= float(input())
    print("circumference = " + str(round(a*pi,2)))
    print("Area = " + str(round(((a/2)**2)*pi,2)))
```

當直徑為 `10` 時，`circle()` 方法應該印出以下結果：<br></br><br></br>
`Circumference = 31.40`<br></br>
`Area = 78.50`<br></br><br></br>
請執行以下程式碼確認結果。

```python
circle()
```

<!-- grade_id: correct_circle, points: 30, grade: true, solution: false -->
```python
"""檢查 circle() 結果是否與以下測試案例相同"""
import unittest
from unittest import mock
from unittest.mock import patch
from io import StringIO

class TestCircle(unittest.TestCase):
    
    d = 6
    
    @patch('builtins.input', side_effect=[d])
    def test_circle(self, mock_input):
        with mock.patch('sys.stdout', new = StringIO()) as output:
            circle()
            self.assertEqual(output.getvalue(), "circumference = 18.84\nArea = 28.26\n")

suite = unittest.TestSuite()
suite.addTest(TestCircle('test_circle'))
result = unittest.TextTestRunner(verbosity=2).run(suite)

assert result.wasSuccessful() == True
```