```python
# 請輸入姓名以及學號，以證名為本人撰寫
Name = ""
NID = ""
```

## 計算 BMI (40 points)

請將下面 `bmi()` 方法完成，輸入身高以及體重並計算 $bmi$ 數值，計算公式是以體重（公斤）除以身高（公尺）的平方。<br></br><br></br>
注意：身高以及體重為 `float` 型態

<!-- grade_id: bmi, points: 0, grade: false, solution: true -->
```python
def bmi():
    """計算 BMI 數值，輸入身高及體重，並印出 BMI 數值"""""
    
    # YOUR CODE HERE
    height=float(input())/100
    weight=float(input())
    print("BMI = {:.2f}".format(weight/height**2))
```

當身高為 175 體重為 74.5 時，`bmi()` 方法應該印出 `BMI = 24.33`，請執行以下程式碼確認結果。

```python
bmi()
```

<!-- grade_id: correct_bmi, points: 40, grade: true, solution: false -->
```python
"""檢查 bmi() 結果是否與以下測試案例相同"""
import unittest
from unittest import mock
from unittest.mock import patch
from io import StringIO

class TestBMI(unittest.TestCase):
    
    hight_input = 170
    weight_input = 66.3
    
    @patch('builtins.input', side_effect=[hight_input, weight_input])
    def test_bmi(self, mock_input):
        with mock.patch('sys.stdout', new = StringIO()) as output:
            bmi()
            self.assertEqual(output.getvalue(), "BMI = 22.94\n")

suite = unittest.TestSuite()
suite.addTest(TestBMI('test_bmi'))
result = unittest.TextTestRunner(verbosity=2).run(suite)

assert result.wasSuccessful() == True
```