```python
# 請輸入姓名以及學號，以證名為本人撰寫
Name = ""
NID = ""
```

## 十二星座 (25 points)

請將下面 `constellation()` 方法完成，輸入月份以及日期並回傳 **星座**。<br></br><br></br>
十二星座如下：
* 水瓶座 ( 1/20 - 2/18 )：Aquarius
* 雙魚座 ( 2/19 - 3/20 )：Pisces
* 牧羊座 ( 3/21 - 4/20 )：Aries
* 金牛座 ( 4/21 - 5/20 )：Taurus
* 雙子座 ( 5/21 - 6/21 )：Gemini
* 巨蟹座 ( 6/22 - 7/22 )：Cancer
* 獅子座 ( 7/23 - 8/22 )：Leo
* 處女座 ( 8/23 - 9/22 )：Virgo
* 天秤座 ( 9/23 - 10/23 )：Libra
* 天蠍座 ( 10/24 - 11/22 )：Scorpio
* 射手座 ( 11/23 - 12/21 )：Sagittarius
* 摩羯座 ( 12/22 - 1/20 )：Capricorn

<!-- grade_id: constellation, points: 0, grade: false, solution: true -->
```python
def constellation():
    """輸入月份及日期，並印出為哪一個星座"""""
    
    # YOUR CODE HERE
    import datetime
    a=int(input())
    b=int(input())
    
    date=datetime.date(2019,a,b)
    
    aries1=datetime.date(2019,3,21)
    aries2=datetime.date(2019,4,20)
    taurus1=datetime.date(2019,4,21)
    taurus2=datetime.date(2019,5,20)
    gemini1=datetime.date(2019,5,21)
    gemini2=datetime.date(2019,6,20)
    cancer1=datetime.date(2019,6,21)
    cancer2=datetime.date(2019,7,22)
    leo1=datetime.date(2019,7,23)
    leo2=datetime.date(2019,8,22)
    virgo1=datetime.date(2019,8,23)
    virgo2=datetime.date(2019,9,22)
    libra1=datetime.date(2019,9,23)
    libra2=datetime.date(2019,10,22)
    scorpio1=datetime.date(2019,10,23)
    scorpio2=datetime.date(2019,11,22)
    sagittarius1=datetime.date(2019,11,23)
    sagittarius2=datetime.date(2019,12,22)
    capricorn1=datetime.date(2019,12,23)
    capricorn2=datetime.date(2019,1,21)
    aquarius1=datetime.date(2019,1,22)
    aquarius2=datetime.date(2019,2,19)
    pisces1=datetime.date(2019,2,20)
    pisces2=datetime.date(2019,3,20)
    
    if(date>=aries1 and date<=aries2):
        print ("Aries")
    elif(date>=taurus1 and date<=taurus2):
        print ("Taurus")
    elif(date>=gemini1 and date<=gemini2):
        print ("Gemini")
    elif(date>=cancer1 and date<=cancer2):
        print ("Cancer")
    elif(date>=leo1 and date<=leo2):
        print ("Leo")
    elif(date>=virgo1 and date<=virgo2):
        print ("Virgo")
    elif(date>=libra1 and date<=libra2):
        print ("Libra")
    elif(date>=scorpio1 and date<=scorpio2):
        print ("Scorpio")
    elif(date>=sagittarius1 and date<=sagittarius2):
        print ("Sagittarius")
    elif(date>=capricorn1 and date<=capricorn2):
        print ("Capricorn")
    elif(date>=aquarius1 and date<=aquarius2):
        print ("Aquarius")
    elif(date>=pisces1 and date<=pisces2):
        print ("Pisces")
```

當月份為 10 日期為 5 時，`constellation()` 方法應該印出 `Libra`，請執行以下程式碼確認結果。

```python
constellation()
```

<!-- grade_id: test_constellation, points: 25, grade: true, solution: false -->
```python
"""檢查 constellation() 結果是否與以下測試案例相同"""
import unittest
from unittest import mock
from unittest.mock import patch
from io import StringIO

class TestConstellation(unittest.TestCase):
    
    testcase_1 = [3, 7]
    testcase_2 = [5, 4]
    
    @patch('builtins.input', side_effect=[testcase_1[0], testcase_1[1]])
    def test_constellation_1(self, mock_input):
        with mock.patch('sys.stdout', new = StringIO()) as output:
            constellation()
            self.assertEqual(output.getvalue(), "Pisces\n")
            
    @patch('builtins.input', side_effect=[testcase_2[0], testcase_2[1]])
    def test_constellation_2(self, mock_input):
        with mock.patch('sys.stdout', new = StringIO()) as output:
            constellation()
            self.assertEqual(output.getvalue(), "Taurus\n")

suite = unittest.TestSuite()
suite.addTest(TestConstellation('test_constellation_1'))
suite.addTest(TestConstellation('test_constellation_2'))
result = unittest.TextTestRunner(verbosity=2).run(suite)

assert result.wasSuccessful() == True
```