```python
# 請輸入姓名以及學號，以證名為本人撰寫
Name = ""
NID = ""
```

## 十二星座 (25 points)

請將下面 `constellation()` 方法完成，輸入月份以及日期並回傳 **星座**。<br></br><br></br>
十二星座如下：
* 水瓶座 ( 1/20 - 2/18 )：Aquarius
* 雙魚座 ( 2/19 - 3/20 )：Pisces
* 牧羊座 ( 3/21 - 4/20 )：Aries
* 金牛座 ( 4/21 - 5/20 )：Taurus
* 雙子座 ( 5/21 - 6/21 )：Gemini
* 巨蟹座 ( 6/22 - 7/22 )：Cancer
* 獅子座 ( 7/23 - 8/22 )：Leo
* 處女座 ( 8/23 - 9/22 )：Virgo
* 天秤座 ( 9/23 - 10/23 )：Libra
* 天蠍座 ( 10/24 - 11/22 )：Scorpio
* 射手座 ( 11/23 - 12/21 )：Sagittarius
* 摩羯座 ( 12/22 - 1/20 )：Capricorn

<!-- grade_id: constellation, points: 0, grade: false, solution: true -->
```python
def constellation():
    """輸入月份及日期，並印出為哪一個星座"""""
    
    # YOUR CODE HERE
    while True:
        month = input()
        date = input()
        digitmonth = month.isdigit()
        digitdate = date.isdigit()
        
        if not (digitmonth and digitdate):
            print("number!")
            continue
        
        month = int(month)
        date = int(date)
        if not (0 < date and date < 32):
            continue
        if not (0 < month and month < 13):
            continue
        
        a = (date <=19 and date >= 1)
        b = (date <= 29 and date >= 24)
        c = (date == 30)
        d = (date == 31)
        e = (date == 20)
        f = (date == 21)
        g = (date == 22)
        h = (date == 23)  
    #一月
            if month == 1: 
            if (a == True):
                print("Capricorn")
                break
            elif b == True:
                print("Aquarius")
                break
            elif e or f:
                print("Capricorn")
                break
            elif g or h:
                print("Aquarius")
                break
            else:
                print("Aquarius")
                break
    #二月
            if month == 2: 
            if a == True:
                print("Aquarius")
                break
            elif b == True:
                print("Pisces")
                break
            elif e or f or g or h:
                print("Pisces")
                break
            else:
                continue
    #三月
            if month == 3: 
            if a == True:
                print("Pisces")
                break
            elif b == True:
                print("Aries")
                break
            elif e == True:
                print("Pisces")
                break
            elif f or g or h:
                print("Ariess")
                break
            else:
                print("Aries")
                break
    #四月
            if month == 4:
            if d == True:
                continue
            if a == True:
                print("Aries")
                break
            elif b == True:
                print("Taurus")
                break
            elif e == True:
                print("Aries")
                break
            elif f or g or h:
                print("Taurus")
                break
            else:
                print("Taurus")
                break
    #五月
            if month == 5: 
            if a == True:
                print("Taurus")
                break
            elif b == True:
                print("Gemini")
                break
            elif e == True:
                print("Taurus")
                break
            elif f or g or h:
                print("Gemini")
                break
            else:
                print("Gemini")
                break
    #六月
            if month == 6: 
            if d ==True:
                continue
            if a == True:
                print("Gemini")
                break
            elif b == True:
                print("Cancer")
                break
            elif e == True:
                print("Gemini")
                break
            elif f or g or h:
                print("Cancer")
                break
            else:
                print("Cancer")
                break
    #七月
            if month == 7: 
            if (a == True):
                print("Cancer")
                break
            elif b == True:
                print("Leo")
                break
            elif e or f or g:
                print("Cancer")
                break
            elif h == True:
                print("Leo")
                break
            else:
                print("Leo")
                break
    #八月
            if month == 8: 
            if (a == True):
                print("Leo")
                break
            elif b == True:
                print("Virgo")
                break
            elif e or f or g:
                print("Leo")
                break
            elif h == True:
                print("Virgo")
                break
            else:
                print("Virgo")
                break
    #九月
            if month == 9: 
            if d == True:
                continue
            if (a == True):
                print("Virgo")
                break
            elif b == True:
                print("Libra")
                break
            elif e or f or g:
                print("Virgo")
                break
            elif h == True:
                print("Libra")
                break
            else:
                print("Libra")
                break
    #十月
            if month == 10: 
            if (a == True):
                print("Libra")
                break
            elif b == True:
                print("Scorpio")
                break
            elif e or f or g:
                print("Libra")
                break
            elif h == True:
                print("Scorpio")
                break
            else:
                print("Scorpio")
                break
    #十一月
            if month == 11: 
            if d == True:
                continue
            if (a == True):
                print("Scorpio")
                break
            elif b == True:
                print("Sagittarius")
                break
            elif e or f or g:
                print("Scorpio")
                break
            elif h == True:
                print("Sagittarius")
                break
            else:
                print("Sagittarius")
                break
    #十二月
            if month == 12: 
            if (a == True):
                print("Sagittarius")
                break
            elif b == True:
                print("Capricorn")
                break
            elif e or f or g:
                print("Sagittarius")
                break
            elif h == True:
                print("Capricorn")
                break
            else:
                print("Capricorn")
                break
```

Output:
```
縮排以及空格：
https://python-book-test.readthedocs.io/en/latest/intro/03_base.html#id3
```

當月份為 10 日期為 5 時，`constellation()` 方法應該印出 `Libra`，請執行以下程式碼確認結果。

```python
constellation()
```

<!-- grade_id: test_constellation, points: 25, grade: true, solution: false -->
```python
"""檢查 constellation() 結果是否與以下測試案例相同"""
import unittest
from unittest import mock
from unittest.mock import patch
from io import StringIO

class TestConstellation(unittest.TestCase):
    
    testcase_1 = [3, 7]
    testcase_2 = [5, 4]
    
    @patch('builtins.input', side_effect=[testcase_1[0], testcase_1[1]])
    def test_constellation_1(self, mock_input):
        with mock.patch('sys.stdout', new = StringIO()) as output:
            constellation()
            self.assertEqual(output.getvalue(), "Pisces\n")
            
    @patch('builtins.input', side_effect=[testcase_2[0], testcase_2[1]])
    def test_constellation_2(self, mock_input):
        with mock.patch('sys.stdout', new = StringIO()) as output:
            constellation()
            self.assertEqual(output.getvalue(), "Taurus\n")

suite = unittest.TestSuite()
suite.addTest(TestConstellation('test_constellation_1'))
suite.addTest(TestConstellation('test_constellation_2'))
result = unittest.TextTestRunner(verbosity=2).run(suite)

assert result.wasSuccessful() == True
```