```python
# 請輸入姓名以及學號，以證名為本人撰寫
Name = ""
NID = ""
```

## 倍數總和 (25 points)
請將下面 `multiple_sum()` 方法完成。

<!-- grade_id: multiple_sum, points: 0, grade: false, solution: true -->
```python
def multiple_sum():
    """輸入兩個整數並計算倍數總和"""""
    
    # YOUR CODE HERE
    a=int(input())
    b=int(input())
    i = 1
    sum=0
    for i in range(i,a+1):
        if i%b == 0:
               sum += i
                print('sum = ' +str(sum))
```

Output:
```
縮排以及空格：
https://python-book-test.readthedocs.io/en/latest/intro/03_base.html#id3
```

當兩個整數分別為 10 體重為 2 時，`multiple_sum()` 方法應該印出 `sum = 30`，請執行以下程式碼確認結果。

```python
multiple_sum()
```

<!-- grade_id: test_multiple_sum, points: 25, grade: true, solution: false -->
```python
"""檢查 gcd() 結果是否與以下測試案例相同"""
import unittest
from unittest import mock
from unittest.mock import patch
from io import StringIO

class TestMultipleSum(unittest.TestCase):
    
    testcase1 = [1000, 3]
    testcase2 = [100, 1]
    testcase3 = [0, 7]
    
    @patch('builtins.input', side_effect=[testcase1[0], testcase1[1]])
    def test_multiple_sum_1(self, mock_input):
        with mock.patch('sys.stdout', new = StringIO()) as output:
            multiple_sum()
            self.assertEqual(output.getvalue(), "sum = 166833\n")
            
    @patch('builtins.input', side_effect=[testcase2[0], testcase2[1]])
    def test_multiple_sum_2(self, mock_input):
        with mock.patch('sys.stdout', new = StringIO()) as output:
            multiple_sum()
            self.assertEqual(output.getvalue(), "sum = 5050\n")
    
    @patch('builtins.input', side_effect=[testcase3[0], testcase3[1]])
    def test_multiple_sum_3(self, mock_input):
        with mock.patch('sys.stdout', new = StringIO()) as output:
            multiple_sum()
            self.assertEqual(output.getvalue(), "sum = 0\n")

suite = unittest.TestSuite()
suite.addTest(TestMultipleSum('test_multiple_sum_1'))
suite.addTest(TestMultipleSum('test_multiple_sum_2'))
suite.addTest(TestMultipleSum('test_multiple_sum_3'))
result = unittest.TextTestRunner(verbosity=2).run(suite)

assert result.wasSuccessful() == True
```