```python
# 請輸入姓名以及學號，以證名為本人撰寫
Name = ""
NID = ""
```

## 最大公因數 (25 points)

請將下面 `gcd()` 方法完成，輸入兩個整數並計算兩個整數的 **最大公因數**。<br></br><br></br>
注意：兩個整數均為 `int` 型態

<!-- grade_id: gcd, points: 0, grade: false, solution: true -->
```python
def gcd():
    """輸入兩個整數並計算兩個整數的最大公因數"""""
    
    # YOUR CODE HERE
    a = int(input())
    b = int(input())
    c = 1
    d = 1
    while d != 0:
        if a<b:
            a,b = b,a
        d = a % b
        c = b
        a,b = b,d
        
    print("GCD =",c)
```

當兩個整數分別為 65 體重為 104 時，`gcd()` 方法應該印出 `GCD = 13`，請執行以下程式碼確認結果。

```python
gcd()
```

<!-- grade_id: test_gcd, points: 25, grade: true, solution: false -->
```python
"""檢查 gcd() 結果是否與以下測試案例相同"""
import unittest
from unittest import mock
from unittest.mock import patch
from io import StringIO

class TestGCD(unittest.TestCase):
    
    testcase1 = [15, 9]
    testcase2 = [65, 104]
    
    @patch('builtins.input', side_effect=[testcase1[0], testcase1[1]])
    def test_gcd_1(self, mock_input):
        with mock.patch('sys.stdout', new = StringIO()) as output:
            gcd()
            self.assertEqual(output.getvalue(), "GCD = 3\n")
            
    @patch('builtins.input', side_effect=[testcase2[0], testcase2[1]])
    def test_gcd_2(self, mock_input):
        with mock.patch('sys.stdout', new = StringIO()) as output:
            gcd()
            self.assertEqual(output.getvalue(), "GCD = 13\n")

suite = unittest.TestSuite()
suite.addTest(TestGCD('test_gcd_1'))
suite.addTest(TestGCD('test_gcd_2'))
result = unittest.TextTestRunner(verbosity=2).run(suite)

assert result.wasSuccessful() == True
```