```python
# 請輸入姓名以及學號，以證名為本人撰寫
Name = ""
NID = ""
```

## 三角形 (25 points)

請將下面 `triangle()` 方法完成，輸入三角形的層數並印出 **左半邊三角形的形狀**。<br></br><br></br>
注意：層數為 `int` 型態

<!-- grade_id: triangle, points: 0, grade: false, solution: true -->
```python
def triangle():
    """輸入三角形的層數並印出左半邊三角形的形狀"""""
    
    # YOUR CODE HERE
    a=int(input())
    for i in range(1,a+1):
        for j in range(1,a+1):
        if i+j<=a:
            print(' ',end=' ')
        elif a==j:
        print('*',end='')
        else:
            print('*',end=' ')
        print()
```

Output:
```
縮排以及空格：
https://python-book-test.readthedocs.io/en/latest/intro/03_base.html#id3
```

當層數為 8 時，`triangle()` 方法應該印出如下的三角形，請執行以下程式碼確認結果。

                  *
                * *
              * * *
            * * * *
          * * * * *
        * * * * * *
      * * * * * * *
    * * * * * * * *

```python
triangle()
```

<!-- grade_id: test_triangle, points: 25, grade: true, solution: false -->
```python
"""檢查 triangle() 結果是否與以下測試案例相同"""
import unittest
from unittest import mock
from unittest.mock import patch
from io import StringIO

class TestTriangle(unittest.TestCase):
    
    testcase1 = 4
    testcase2 = 6
    testcase3 = 8
    
    @patch('builtins.input', side_effect=[testcase1])
    def test_triangle_1(self, mock_input):
        with mock.patch('sys.stdout', new = StringIO()) as output:
            triangle()
            self.assertEqual(output.getvalue(), "      *\n    * *\n  * * *\n* * * *\n")
            
    @patch('builtins.input', side_effect=[testcase2])
    def test_triangle_2(self, mock_input):
        with mock.patch('sys.stdout', new = StringIO()) as output:
            triangle()
            self.assertEqual(output.getvalue(), "          *\n        * *\n      * * *\n    * * * *\n  * * * * *\n* * * * * *\n")
            
    @patch('builtins.input', side_effect=[testcase3])
    def test_triangle_3(self, mock_input):
        with mock.patch('sys.stdout', new = StringIO()) as output:
            triangle()
            self.assertEqual(output.getvalue(), "              *\n            * *\n          * * *\n        * * * *\n      * * * * *\n    * * * * * *\n  * * * * * * *\n* * * * * * * *\n")

suite = unittest.TestSuite()
suite.addTest(TestTriangle('test_triangle_1'))
suite.addTest(TestTriangle('test_triangle_2'))
suite.addTest(TestTriangle('test_triangle_3'))
result = unittest.TextTestRunner(verbosity=2).run(suite)

assert result.wasSuccessful() == True
```