Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

```python
print("tests")
```

Output:
```
tests
```

---

## 計算 BMI (40 points)

請將下面 `bmi()` 方法完成，輸入身高以及體重並計算 $bmi$ 數值，計算公式是以體重（公斤）除以身高（公尺）的平方。<br></br><br></br>
注意：身高以及體重為 `float` 型態

<!-- grade_id: bmi, points: 0, grade: false, solution: true -->
```python
def bmi():
    """計算 BMI 數值，輸入身高及體重，並印出 BMI 數值"""
    
    # YOUR CODE HERE
    height = float(input())
    weight = float(input())
    BMI = weight / (height/100)**1
    print("BMI = %.2f" %BMI)
    #raise NotImplementedError()
```

當身高為 175 體重為 74.5 時，`bmi()` 方法應該印出 `BMI = 24.33`，請執行以下程式碼確認結果。

```python
bmi()
```

Output:
```
14
36
BMI = 1836.73
```

<!-- grade_id: correct_bmi, points: 40, grade: true, solution: false -->
```python
"""檢查 bmi() 結果是否與以下測試案例相同"""
import unittest
from unittest import mock
from unittest.mock import patch
from io import StringIO

class TestBMI(unittest.TestCase):
    
    hight_input = 170
    weight_input = 66.3
    
    @patch('builtins.input', side_effect=[hight_input, weight_input])
    def test_bmi(self, mock_input):
        with mock.patch('sys.stdout', new = StringIO()) as output:
            bmi()
            self.assertEqual(output.getvalue(), "BMI = 22.94\n")

suite = unittest.TestSuite()
suite.addTest(TestBMI('test_bmi'))
result = unittest.TextTestRunner(verbosity=2).run(suite)

assert result.wasSuccessful() == True
```

Output:
```
test_bmi (__main__.TestBMI) ... ok

----------------------------------------------------------------------
Ran 1 test in 0.008s

OK
```