```python
# 請輸入姓名以及學號，以證名為本人撰寫
Name = ""
NID = ""
```

## 十二星座 (25 points)

請將下面 `constellation()` 方法完成，輸入月份以及日期並回傳 **星座**。<br></br><br></br>
十二星座如下：
* 水瓶座 ( 1/20 - 2/18 )：Aquarius
* 雙魚座 ( 2/19 - 3/20 )：Pisces
* 牧羊座 ( 3/21 - 4/20 )：Aries
* 金牛座 ( 4/21 - 5/20 )：Taurus
* 雙子座 ( 5/21 - 6/21 )：Gemini
* 巨蟹座 ( 6/22 - 7/22 )：Cancer
* 獅子座 ( 7/23 - 8/22 )：Leo
* 處女座 ( 8/23 - 9/22 )：Virgo
* 天秤座 ( 9/23 - 10/23 )：Libra
* 天蠍座 ( 10/24 - 11/22 )：Scorpio
* 射手座 ( 11/23 - 12/21 )：Sagittarius
* 摩羯座 ( 12/22 - 1/20 )：Capricorn

<!-- grade_id: constellation, points: 0, grade: false, solution: true -->
```python
def constellation():
    """輸入月份及日期，並印出為哪一個星座"""""
    
    # YOUR CODE HERE
    month = input("")
    date = input("")
    month = int(month)
    date = int(date)
    def March(date):
        if( date>=1 and date<=20):
            return "Pisces"
        elif( date>=21 and date<=31 ):
            return "Aries"
    def April(date):
        if( date>=1 and date<=20):
            return "Aries"
        elif( date>=21 and date<=30 ):
            return "Taurus"
    def May(date):
        if( date>=1 and date<=20):
            return "Taurus"
        elif( date>=21 and date<=31 ):
            return "Gemini"
    def June(date):
        if( date>=1 and date<=20):
            return "Gemini"
        elif( date>=21 and date<=30 ):
            return "Cancer"
    def July(date):
        if( date>=1 and date<=22):
            return "Cancer"
        elif( date>=23 and date<=31 ):
            return "Leo"
    def Auguest(date):
        if( date>=1 and date<=22):
            return "Leo"
        elif( date>=23 and date<=31 ):
            return "Virgo"
    def Sebtember(date):
        if( date>=1 and date<=22):
            return "Virgo"
        elif( date>=23 and date<=30 ):
            return "Libra"
    def October(date):
        if( date>=1 and date<=22):
            return "Libra"
        elif( date>=23 and date<=31 ):
            return "Scorpio"
    def November(date):
        if( date>=1 and date<=22):
            return "Scorpio"
        elif( date>=23 and date<=30 ):
            return "Sagittarius"
    def December(date):
        if( date>=1 and date<=22):
            return "Sagittarius"
        elif( date>=23 and date<=31 ):
            return "Capricorn"
    def January(date):
        if( date>=1 and date<=21):
            return "Capricorn"
        elif( date>=22 and date<=31 ):
            return "Aquarius"
    def February(date):
        if( date>=1 and date<=19):
            return "Aquarius"
        elif( date>=20 and date<=29 ):
            return "Pisces"
        
    monthes = {
                    3 : March(date),
            4 : April(date),
            5 : May(date),
            6 : June(date),
            7 : July(date),
            8 : Auguest(date),
            9 : Sebtember(date),
            10 : October(date),
            11 : November(date),
            12 : December(date),
            1 : January(date),
            2 : February(date),
    }
    print(monthes[month])
    }
```

Output:
```
File "<ipython-input-1-ef51d9580711>", line 85
    }
    ^
SyntaxError: invalid syntax

語法錯誤，可能缺少括號、字符等等，請檢查語法是否正確。
```

當月份為 10 日期為 5 時，`constellation()` 方法應該印出 `Libra`，請執行以下程式碼確認結果。

```python
constellation()
```

<!-- grade_id: test_constellation, points: 25, grade: true, solution: false -->
```python
"""檢查 constellation() 結果是否與以下測試案例相同"""
import unittest
from unittest import mock
from unittest.mock import patch
from io import StringIO

class TestConstellation(unittest.TestCase):
    
    testcase_1 = [3, 7]
    testcase_2 = [5, 4]
    
    @patch('builtins.input', side_effect=[testcase_1[0], testcase_1[1]])
    def test_constellation_1(self, mock_input):
        with mock.patch('sys.stdout', new = StringIO()) as output:
            constellation()
            self.assertEqual(output.getvalue(), "Pisces\n")
            
    @patch('builtins.input', side_effect=[testcase_2[0], testcase_2[1]])
    def test_constellation_2(self, mock_input):
        with mock.patch('sys.stdout', new = StringIO()) as output:
            constellation()
            self.assertEqual(output.getvalue(), "Taurus\n")

suite = unittest.TestSuite()
suite.addTest(TestConstellation('test_constellation_1'))
suite.addTest(TestConstellation('test_constellation_2'))
result = unittest.TextTestRunner(verbosity=2).run(suite)

assert result.wasSuccessful() == True
```