```python
# 請輸入姓名以及學號，以證名為本人撰寫
Name = ""
NID = ""
```

## 十二星座 (25 points)

請將下面 `constellation()` 方法完成，輸入月份以及日期並回傳 **星座**。<br></br><br></br>
十二星座如下：
* 水瓶座 ( 1/20 - 2/18 )：Aquarius
* 雙魚座 ( 2/19 - 3/20 )：Pisces
* 牧羊座 ( 3/21 - 4/20 )：Aries
* 金牛座 ( 4/21 - 5/20 )：Taurus
* 雙子座 ( 5/21 - 6/21 )：Gemini
* 巨蟹座 ( 6/22 - 7/22 )：Cancer
* 獅子座 ( 7/23 - 8/22 )：Leo
* 處女座 ( 8/23 - 9/22 )：Virgo
* 天秤座 ( 9/23 - 10/23 )：Libra
* 天蠍座 ( 10/24 - 11/22 )：Scorpio
* 射手座 ( 11/23 - 12/21 )：Sagittarius
* 摩羯座 ( 12/22 - 1/20 )：Capricorn

<!-- grade_id: constellation, points: 0, grade: false, solution: true -->
```python
def constellation():
    """輸入月份及日期，並印出為哪一個星座"""""
    
    # YOUR CODE HERE
    m=int(input())
    d=int(input())
    
    if(m==1 and(d>=22 and d<31))or(m==2 and(d<=19 and d>=0)):
        print("Aquarius")
    elif(m==2 and(d>=20 and d<29))or(m==3 and(d<=20 and d>=0)):
        print("Pisces")
    elif(m==3 and(d>=21 and d<31))or(m==4 and(d<=20 and d>=0)):
        print("Aries")
    elif(m==4 and(d>=21 and d<30))or(m==5 and(d<=20 and d>=0)):
        print("Taurus")
    elif(m==5 and(d>=21 and d<31))or(m==6 and(d<=20 and d>=0)):
        print("Gemini")
    elif(m==6 and(d>=21 and d<30))or(m==7 and(d<=22 and d>=0)):
        print("Cancer")
    elif(m==7 and(d>=23 and d<31))or(m==8 and(d<=22 and d>=0)):
        print("Leo")
    elif(m==8 and(d>=23 and d<31))or(m==9 and(d<=22 and d>=0)):
        print("Virgo")
    elif(m==9 and(d>=23 and d<30))or(m==10 and(d<=22 and d>=0)):
        print("Libra")
    elif(m==10 and(d>=23 and d<31))or(m==11 and(d<=22 and d>=0)):
        print("Scorpio")
    elif(m==11 and(d>=23 and d<30))or(m==12 and(d<=22 and d>=0)):
        print("Sagittarius")
    else:
        print("Capricorn")
```

當月份為 10 日期為 5 時，`constellation()` 方法應該印出 `Libra`，請執行以下程式碼確認結果。

```python
constellation()
```

Output:
```
1
15
Capricorn
```

<!-- grade_id: test_constellation, points: 25, grade: true, solution: false -->
```python
"""檢查 constellation() 結果是否與以下測試案例相同"""
import unittest
from unittest import mock
from unittest.mock import patch
from io import StringIO

class TestConstellation(unittest.TestCase):
    
    testcase_1 = [3, 7]
    testcase_2 = [5, 4]
    
    @patch('builtins.input', side_effect=[testcase_1[0], testcase_1[1]])
    def test_constellation_1(self, mock_input):
        with mock.patch('sys.stdout', new = StringIO()) as output:
            constellation()
            self.assertEqual(output.getvalue(), "Pisces\n")
            
    @patch('builtins.input', side_effect=[testcase_2[0], testcase_2[1]])
    def test_constellation_2(self, mock_input):
        with mock.patch('sys.stdout', new = StringIO()) as output:
            constellation()
            self.assertEqual(output.getvalue(), "Taurus\n")

suite = unittest.TestSuite()
suite.addTest(TestConstellation('test_constellation_1'))
suite.addTest(TestConstellation('test_constellation_2'))
result = unittest.TextTestRunner(verbosity=2).run(suite)

assert result.wasSuccessful() == True
```

Output:
```
test_constellation_1 (__main__.TestConstellation) ... ok
test_constellation_2 (__main__.TestConstellation) ... ok

----------------------------------------------------------------------
Ran 2 tests in 0.007s

OK
```